---
title: Test
description: Test

---
## Sample exercise

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: 89ce975940
```


`@instructions`
Print the first value in the array 
`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
import numpy as np
a = np.arange(20)

print(a[...])
```

`@solution`
```{python}
a[0]
```

`@sct`
```{python}

```
